# Coding Lab 2: Variables, Operators, Functions


Welcome to the second coding lab!

CodingLab labs are meant to be interactive - so you are encouraged discuss with others, but each still fill out your own notebook for submission.

CodingLabs are also meant to be exploratory. There are broad questions in the notebook that you should explore, and try to answer - but you are also very much encouraged to explore other related ideas as you go! If you explore something beyond what is required in the notebook, **feel free to leave that "extra" code and any notes to yourself in the notebook you submit.**

If you have a question about how something works / what something does - try it out, and see what happens! If you can't figure it out, reach out to your instructional staff - we're here to help!

## Part 1: Variables

Variables are used in Python to store values. 

### Defining Variables

Declare an int, float, string, and boolean. Fill them with with any values you want. 

Call them `my_int`, `my_float`, `my_string` and `my_boolean`. 

<!-- grade_id: cell-bc577641f19eb639, points: 0, grade: false, solution: true -->
```python
# your code here
```

You'll see a lot of `assert` test cells throughout this course - most often in exams and assignments, but sometimes in coding labs. These are meant to ensure you're on the right track!

When you see `assert` followed by a variable name, that is Python checking to ensure that variable exists. If you mistyped the variable name above (when defining the variable name), you would get an assertion error when running the cell below. The assertion error will point to which `assert` led the error to be raised. It's then your job to go back to your code and figure out why the `assert` test failed!

If you're on the right track, there will be no output - or as you'll hear us say, these will "pass siltently". **This doesn't guarantee your answer is right...but it let's you know you're on the right track!**

<!-- grade_id: cell-74f3c0118c73cedc, points: 0, grade: true, solution: false -->
```python
assert my_int
assert my_float
assert my_string
assert my_boolean
```

Use the following cells to check the types of the variables you write are as expected

```python
type(my_int)
```

```python
type(my_float)
```

```python
type(my_string)
```

```python
type(my_boolean)
```

Note that we could also use `assert` statements to test the type. If the `type` of `my_int` is an integer (`int`), this cell will "pass silently" (meaning give no output). If it raised an asserttion error, we would know we had done something wrong above.

<!-- grade_id: cell-0648f3aeaf31e974, points: 0, grade: true, solution: false -->
```python
# test that my_int is an int
assert type(my_int) == int
```

## Part 2: Operators & Comparisons

Operators are used in Python to do things with variables. 

### Operator Questions

First - explore using Python operators and print to see what is the result. Make sure you can add and subtract numbers, concatenate strings, do boolean comparisons, etc. For example, `print(1+1)`.

<!-- grade_id: cell-8f0d4b798a73fcc8, points: 0, grade: false, solution: true -->
```python
## Do some exploring!

# your code here
```

Once you have explored using operators, try to answer some quesions using them. 

Let's start with an example: is the multiplication of 13 and 15 greater than 200?

```python
13 * 15 > 200
```

Similar to that example, try to answer the following questions using Python variables & comparisons. You can assign the output to whatever variable name you like or print them:
- Is the remainder of dividing 323 by 13 greater than 6?
- Is 7 to the power of 3 greater than or equal to 300?
- Is 49 squared greater than 2500 or (boolean comparison) is 14 to the power of 3 greater than 2500
- Is 49 modulus 9 equal to 67 modulus 9 and (boolean comparison) is the sum 49 & 67 modulus 9 greater than 7
- Bonus: Set the variable `comp_1` to the result of whether 17 squared is not the same as 867 divided by 3 (it can be written in "at least" two ways.)

<!-- grade_id: cell-14b037c8bfb689e5, points: 0, grade: false, solution: true -->
```python
# your code here
```

## Part 3: Functions

Functions are a way to organize code into a procedure that we can repeat and use whenever we want. 

### Function Questions

#### Write a function that takes two inputs, adds them together, and returns the result

Call this function `add_two`.

<!-- grade_id: cell-bc8d2072710911df, points: 0, grade: false, solution: true -->
```python
def add_two(_FILL_IN_INPUTS_):
    
# your code here
```

```python
# Check that you can use the function
add_two(2, 2)
```

<!-- grade_id: cell-20f5706f33ded922, points: 0, grade: true, solution: false -->
```python
# Check this function returns the correct number
assert add_two(2, 2) == 4
assert add_two(3, 2) == 5
```

#### Write a function 

Write a function `mult_two` that takes one input, multiplies it by two and returns the result

<!-- grade_id: cell-8d0f3557ad318843, points: 0, grade: false, solution: true -->
```python
# your code here
```

```python
# Check that you can use the function
mult_two(2)
```

<!-- grade_id: cell-0c88c4b83c2c9f6b, points: 0, grade: true, solution: false -->
```python
# Check this function returns the correct outputs
assert mult_two(2) == 4
assert mult_two(3) == 6
```

## Part 4: Asserts

You will notice, through the coding labs and assignments, code that looks something like `assert True`, in which the assert statement is used. The goal, in your coding labs and assignments is for `assert` statements to "pass silently"...meaning give no output.

Using assert essentially means to `assert that the following code is True`.

Or, slightly more formally, `assert that the following code evaluates as True`.

If the asserted code does not evaluate as True, it raises an error - meaning it interrupts the code execution because something went wrong.

This can be used as a way to test that code does what you expect it to do, and is therefore part of how we will test your code programmatically. 

Because the use of `assert` will be common in course materials, let's take a moment to explore how it works and what happens when we use it.

In the cell below: declare a variable called `boolean` that  has type boolean, and passes the assert in the next two cells:

<!-- grade_id: cell-132ecd7af4833ad3, points: 0, grade: false, solution: true -->
```python
# your code here
```

```python
# Check that there is variable called `boolean` that is a boolean
assert type(boolean) == bool
```

```python
# Check that the variable called `boolean` passes the assert
assert boolean
```

###  Assert Explorations

Now, explore using assert. 

What happens if you assert a True comparison, like `6 < 10`? a False comparison, like `6 > 10`?

What happens if you assert an integer (check what happens when assert 0)? or a string? 

What happens if you assert None? 

<!-- grade_id: cell-355e2ee2814f7325, points: 0, grade: false, solution: true -->
```python
# your code here
```

## The End!

This is the end of this Coding Lab!

Be sure you've made a concerted effort to complete all the tasks specified in this lab. Then, go ahead and submit on Datahub!