# CL3: Conditionals, Collections

Welcome to the third coding lab!

In this Coding Lab we will focus on working with conditionals, collections and get additional practice with conditionals and functions.

## Part 1: Conditionals

Conditionals a form of control flow for executing certain code if a specific condition is met. 

### Conditional Questions

In this part, we will re-visit some questions about conditionals, and start combining them together with collections. 

Make sure you can use if/elif/else statements to control which code blocks are run.

#### Controlling Output With Conditionals I

Using the code below, keep updating the variable value to make sure you can make the conditional return each of the possibilities. 

```python
# ToDo: Update value to control what gets executed
value = 10

# This code provided
if type(value) == int:
    if value < 10:
        result = 'small value'
    if value >= 10:
        result = 'big value'
else:
    result = 'false value'
```

```python
print(result)
```

#### Controlling Output With Conditionals II

In the cell below, first write something into each `result` assignment string below. For example, it could be something like: "B1 is False & B2 is True". 

Then make sure you can change values of `b1` & `b2` to make the program execute each option. 

```python
# ToDo: Update strings stored in result to control what gets executed
b1 = False
b2 = False

# Next, try this out
if b1 and not b2:
    result = ''
elif not b1 and b2:
    result = ''
elif b1 and b2:
    result = ''
else:
    result = ''
```

```python
print(result)
```

### Conditional Challenge

The following is more challenging question relating to conditional statements. If it seem too tricky for now, just skip ahead to the next section.

#### Conditional Challenge #1

Input: assumes a variable `val_1`, that can be either a string or an integer

Output: based on the conditions, this code block will assign a variable called `output` to be a boolean

Use conditions to check whether `val_1` is a string or an integer object:
- If it is a string, check whether it's length is greater then 8
    - If so, set the value `output` to True
- If it is an integer, compare whether it is less than 100
    - If so, set the value `output` to True
- Else, set `output` to False

<!-- grade_id: cell-2f262d141db7a722, points: 0, grade: false, solution: true -->
```python
# your code here
```

```python
assert type(output) == bool
```

## Part 2: Collections

Collections are Python variable types than can store a 'collection' of items.

Here we'll practice defining collections and indexing.

### Collection Questions

Create a list, a tuple, and a dictionary. Fill them with with any values you want. 

Call them `my_list`, `my_tuple`, and `my_dictionary`.

<!-- grade_id: cell-03b4194c0ccc5cb9, points: 0, grade: false, solution: true -->
```python
# your code here
```

Use the following cells to check the types of the variables you write are as expected

```python
type(my_list)
```

```python
type(my_tuple)
```

```python
type(my_dictionary)
```

On an assignment, you'd likely see the same concept tested with the following `assert` statements

<!-- grade_id: cell-00392a1d77c0cd31, points: 0, grade: true, solution: false -->
```python
assert type(my_list) == list
assert type(my_tuple) == tuple
assert type(my_dictionary) == dict
```

#### Declaring Collections

Note that there can be more than one way to declare collections. 

As well as declaring them with '[]' and '()', we can use the list & tuple constructors, as shown below. 

```python
# Run me! - This creates another list & tuple, using the list & tuple constructors
some_list = list([1, 2, 3, 4, 5])
some_tuple = tuple(["peanut", "butter", "and", "jelly"])
some_dict = dict([('First', 'Shannon'), ('Last', 'Ellis')])

print("'some_list' contains: \t", some_list)
print("'some_tuple' contains: \t", some_tuple)
print("'some_dict' contains: \t", some_dict)
```

Compare the types of these variables to your equivalent variables to confirm.

In the cell below, write in the type of the objects that you expect them to be, and make sure the asserts pass. 

<!-- grade_id: cell-e165e0773228e7df, points: 0, grade: true, solution: true -->
```python
# Fill in `_WRITE_IN_TYPE_HERE` with the type you expect each variable to be
assert type(some_list) == _WRITE_IN_TYPE_HERE_
assert type(some_tuple) == _WRITE_IN_TYPE_HERE_
assert type(some_dict) == _WRITE_IN_TYPE_HERE_
```

### Indexing

Given the following list (`my_lst`, defined below), do the following operations:
- Get the length of the list (assign this to a variable `lst_len`)
- Get the 1st element of the list (call the selection `ind1`)
- Get the last element of the list (call the selection `ind2`)
- Get the second to the fourth element of the list (call the selection `ind3`)
- Get from the fifth element, to the end of the list (call the selection `ind4`)
- Get every second element of the list, starting at the first element (call the selection `ind5`)
- Get every second element of the list, starting at the second element (call the selection `ind6`)

```python
my_lst = ['a', True, 12, 'tomato', False, None, 23, 'python', [], 5.5]
```

<!-- grade_id: cell-09da9aaa2dc855ad, points: 0, grade: false, solution: true -->
```python
# your code here
```

## Part 3: Collections + Conditionals

Collections can work with all of the things we've discussed up to this point - variables, conditionals, functions, etc. Here, we'll get extra practice combining collections with conditionals.

### Collections + Conditionals Question

Imagine you've asked a group of young children to make a pile of at least 5 toys. 

You then want to write code to determine if the kid accomplished the task. Specifically, if the pile has at least 5 toys, you want `output` to store the string `'success'`. Otherwise, you want `output` to store the string `'try again'`

The toys placed in one of these kids' piles is defined below in the list `kid_a`. Using this list, and conditionals, write code to accomplish the task above! 

```python
kid_a = ['truck', 'barbie', 'coloring book', 'dinosaur']
```

<!-- grade_id: cell-fcc40c86d462e69b, points: 0, grade: false, solution: true -->
```python
# your code here
```

<!-- grade_id: cell-79153e133036327d, points: 0, grade: true, solution: false -->
```python
assert output in ['success', 'try again']
```

## Part 4: Collections + Conditionals + Functions

Let's put all of these conepts together...!

### Collections + Conditionals + Functions Question 

The code you generated in Part 3 will only work on a single list `kid_a`. Functions, on the other hand, are ways to define a set of instructions that can operate on various input, returning the correct output.

Here, write a function `toy_pile()` that accomplishes the same task as above, but takes a list of toys as input and `returns` either `'success'` or `'try again'`, depending upon how many toys are in the list of toys.


<!-- grade_id: cell-18d3066acd89631e, points: 0, grade: false, solution: true -->
```python
# your code here
```

```python
# test out your function here 
```

<!-- grade_id: cell-4c017eb2981df277, points: 0, grade: true, solution: false -->
```python
assert toy_pile(kid_a) == 'try again'
```

<!-- grade_id: cell-ee496957763dc6bd, points: 0, grade: true, solution: false -->
```python
assert toy_pile(['lego', 'princess dress', 'tinker toy', 'doll', 'stroller']) == 'success'
```

### Putting it all together

Now that you've seen an example, think up another situation where you could write a function that incorporates the following:

- takes a dictionary as input
- utilizes a conditional within the function
- carries out some operation and `return`s some output

First, explain what your function does in plain English in the cell below. Then, write the python code for the function in the following cell and test it out in the last cell.

Explain your function here (in words, not code)

<!-- grade_id: cell-d6492d162ed04798, points: 0, grade: false, solution: true -->
```python
# define your function here

# your code here
```

```python
# test out/execute your function here
```

## The End!

This is the end of this Coding Lab!

Be sure you've made a concerted effort to complete all the tasks specified in this lab. Then, go ahead and submit on datahub!